Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your github nickname below:

```python
NICKNAME = "Donil0"
```

---

Write a function which solves a quadratic equation, $x^2 + b x + c = 0$. Your function must receive two values for coefficients and return two real or complex roots. Always return two roots, even if the equation has a single double root. 

<!-- grade_id: cell-2032efc64a327acc, points: 0, grade: false, solution: true -->
```python
from math import sqrt
from cmath import sqrt as c_sqrt
from decimal import Decimal
def solve_quad(b, c):
    ''' Parameters
    ----------
    b, c : float
       Coefficients
       
    Returns
    -------
    x1, x2 : float or complex
        Roots.
    '''
    # YOUR CODE HERE         
    d = b ** 2 - 4 * c
    if 4*c/b**2 <1e-15:
        return (-c/b, -b+c/b)
    if d >= 0:
        x2, x1 = (-b + sqrt(d)) / 2, (-b - sqrt(d)) / 2
    else:
        x1, x2 = (-b + c_sqrt(d)) / 2, (-b - c_sqrt(d)) / 2
    return x1, x2
     
```

<!-- grade_id: cell-89cf7badb387b0e0, points: 5, grade: true, solution: false -->
```python
#
# Your implementation should pass tests in this cell.
#
# Do not remove or alter this cell. You may run it, but do not remove or alter it.
# Your changes to this cell will be ignored on grading.
# There will be additional tests.
#
from numpy.testing import assert_allclose

variants = [{'b': 4.0, 'c': 3.0},
            {'b': 2.0, 'c': 1.0},
            {'b': 0.5, 'c': 4.0},
            {'b': 1e10, 'c': 3.0},
           ]

for var in variants:
    x1, x2 = solve_quad(**var)
    assert_allclose(x1*x2, var['c'], rtol=1e-15)
```

Output:
```
[... skipping hidden 1 frame]

File C:\Games\Piska\lib\site-packages\numpy\testing\_private\utils.py:844, in assert_array_compare(comparison, x, y, err_msg, verbose, header, precision, equal_nan, equal_inf)
    840         err_msg += '\n' + '\n'.join(remarks)
    841         msg = build_err_msg([ox, oy], err_msg,
    842                             verbose=verbose, header=header,
    843                             names=('x', 'y'), precision=precision)
--> 844         raise AssertionError(msg)
    845 except ValueError:
    846     import traceback

AssertionError: 
Not equal to tolerance rtol=1e-15, atol=0

Mismatched elements: 1 / 1 (100%)
Max absolute difference: 3.
Max relative difference: 1.
 x: array(-0.)
 y: array(3.)
```

<!-- grade_id: cell-0970e4298cb77486, points: 5, grade: true, solution: false -->
```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```